<center><h3>[CS152] HARNESSING ARTIFICIAL INTELLIGENCE ALGORITHMS</h3></center>
<center><h3>Taha Bouhoun</h3></center>

Before you turn in this assignment, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then run the test cells for each of the questions you have answered.  Note that a grade of 3 for the related LOs requires all tests in the "Basic Functionality" section to be passed.  The test cells pass if they execute with no errors (i.e. all the assertions are passed).

Make sure you fill in any place that says `YOUR CODE HERE`.  Be sure to remove the `raise NotImplementedError()` statements as you implement your code - these are simply there as a reminder if you forget to add code where it's needed.

---

<h1>CS152 Assignment 2: The DPLL Algorithm</h1>

<h1>Question 1</h1>

Define your <code>Literal</code> class below.  Ensure that you define a <code>name</code> and <code>sign</code> attribute as required in the assignment description.  In addition, include any other attributes and member functions as needed. You will need to overload <code>__neg__()</code> to correctly handle negated literals.  You may need to overload <code>__eq__()</code> and <code>__hash__()</code> also.

<!-- grade_id: cell-575ae3fbbe1ccbf6, points: 0, grade: false, solution: true -->
```python
# Import any packages you need here
# Also define any variables as needed

from copy import*

class Literal:
    def __init__(self, name, sign=True):
        self.name = name
        self.sign = sign
        
    def __neg__(self):
        return Literal(self.name, not self.sign)

    def __eq__(self, other):
        return self.name == other.name
    
    def __str__(self):
        return "-" + self.name if not self.sign else self.name
    
    def __hash__(self):
        return hash(self.name)
```

<h1>Question 2</h1>
Implement DPLL by filling in the stubs below.  Note that the various heuristics are not required to be implemented for basic functionality, but a template has been provided for you to do so if you attempt the extension questions

## Pseudo code of DPLL algorithm:
<code>Code by: Ruml Wheeler [YouTube]</code>

DPLL($\alpha$):
- if $\alpha$ has no clauses, return True
- if $\alpha$ has an empty clause, return False
- if $\alpha$ contains a unit clause, return **DPLL(SIMPLIFY($\alpha$, leteral))**
- v $\leftarrow$ choose a variable in $\alpha$
- if **DPLL(SIMPLIFY($\alpha$, v))** is True, return True
    - else, return **DPLL(SIMPLIFY(α, $\neg$v)**



SIMPLIFY($\alpha$, literal):
- remove clauses in $\alpha$ where literal is positive
- remove $\neg$literal form clauses where it appears
- return new $\alpha$

<!-- grade_id: cell-2846b43cfaf88026, points: 0, grade: false, solution: true -->
```python
# Define degree, pure symbol and unit clause heuristics here (optional).
# Uncomment the lines below and add in your functions

# def degree_heuristic(clauses, symbols, model):
"""
  S = degree_heuristic(clauses, symbols, model):
      clauses: The KB with a list of clauses
      symbols: A set of the current literals 
      (we are only interested in the literal, not the sign)
      model: A dictionary giving the truth values assigned for the model
      S: A literal object representing the unassigned 
          symbol that appears the most across all clauses (should be positive)
"""

# def find_pure_symbol(clauses, symbols, model):
"""
  P = find_pure_symbol(clauses, symbols, model):
      clauses: The KB with a list of clauses
      symbols: A set of the current literals 
      (we are only interested in the literal, not the sign)
      model: A dictionary giving the truth values assigned for the model
      P: A pure symbol (currently unassigned in model)
"""

# def find_unit_clause(clauses, symbols, model):
"""
  U = find_unit_clause(clauses, symbols, model):
      clauses: The KB with a list of clauses
      symbols: A set of the current literals 
      (we are only interested in the literal, not the sign)
      model: A dictionary giving the truth values assigned for the model
      U:  A unit clause (currently unassigned in model)
"""

def DPLL_Satisfiable(KB):
    """ 
    satisfiable, model = DPLL_Satisfiable(KB)
    Takes in a KB and returns whether the KB is satisfiable, 
    and the model that makes it so
        
    KB: A knowledge base of clauses (CNF)
    consisting of a list of sets of literals.  
    A KB might look like [{A, B},{-A ,C ,D}]
    
    satisfiable: Returns True if the KB is satisfiable, False otherwise.
    
    Model: A dictionary that assigns a truth value 
    to each literal for the model that satisfies KB.
    """
    
    T = Literal('T')
    # Create a set the literals
    symbols = set()
    # Adds all the literals in the KB to the set of symbols
    [[symbols.add(literal) for literal in clause] for clause in KB]
    
    
    def dpll(clauses, symbols, model):
        # Iterates over the clauses of the KB
        for clause in clauses:
            count = 0
            
            for lit in clause:
                if lit.name in model:
                    count += 1
                    if lit.sign == model[lit.name]:
                        clause.add(T)
                        break                        
                    elif count == len(clause):
                        return False, {}
        
        # If (at least) one literal from each clause is True,
        # then the model is satisfiable.
        if all([T in clause for clause in clauses]):
            return True, model
        
        while symbols:
            # Picks a literal from the symbol set
            pick = symbols.pop()
            model1, model2 = deepcopy(model), deepcopy(model)
            # Associates the two possible truth value to the picked literal
            model1[pick.name], model2[pick.name] = True, False
            
            # Recursive call to test satisfiability given the literal is True
            satisfiable, Model = dpll(deepcopy(clauses),
                                      deepcopy(symbols),
                                      model1)
            if satisfiable:
                return satisfiable, Model
            
            # Recursive call to test satisfiability given the literal is Flase
            satisfiable, Model = dpll(deepcopy(clauses),
                                      deepcopy(symbols),
                                      model2)            
            if satisfiable:
                return satisfiable, Model
        
        # If we exhausted the possible truth value combinations of 
        # all literals without having a model that returns True
        # then we can deduce that the model is unsatisfiable.
        return False, {}
    
    return dpll(KB, symbols, {})
```

<h1>Question 3</h1>

Implement your KB from Russell & Norvig in CNF by filling in the sets inside the list <code>KB</code> below.  Ensure that your KB is in a list of set format as stated in the assignment instructions.

<!-- grade_id: cell-2600ca4ae2649468, points: 0, grade: false, solution: true -->
```python
# KB of exercise 7.20 from Russell & Norvig:

A = Literal('A')
B = Literal('B')
C = Literal('C')
D = Literal('D')
E = Literal('E')
F = Literal('F')

KB = [{-A, B, E},{-B, A},{-E, A},
      {-E, D},
      {-C, -F, -B},
      {-E, B},
      {-B, F},
      {-B, C}]

DPLL_Satisfiable(KB)
```

Output:
```
(True, {'F': True, 'E': False, 'D': True, 'C': True, 'B': False, 'A': False})
```

<h1>Extensions</h1>

1. Complete the degree heuristic function by filling in the placeholder in Q2.
2. Complete the pure symbol and unit clause heuristics by filling in the placeholder in Q2.
3. Modify the pure symbol heuristic to choose the pure symbol that occurs in the most number of clauses.

<h1>Basic Functionality Tests</h1>

All of the tests in this section must be passed for the code to be considered fully functional.  Other unseen test cases will be used after submission to further verify functionality

<!-- grade_id: cell-6af1cbc180046559, points: 1, grade: true, solution: false -->
```python
# This section will test the correct definition of the literal class

# Define some literals
A = Literal('A')
B = Literal('B')
C = Literal('C')
D = Literal('D')
E = Literal('E')
F = Literal('F')


# Test the name attribute works correctly
assert(A.name == 'A')

# Test that negation works correctly
assert(not (-C).sign)
```

<!-- grade_id: cell-76b5f4949b4918a0, points: 2, grade: true, solution: false -->
```python
# This section will test that the KB has been correctly converted to CNF by performing 
# tests over all possible models
import itertools
symbols = {A,B,C,D,E,F}
symbol_list = [A,B,C,D,E,F]

def evaluate_russell_norvig_KB(model):
    # Manually evaluate the KB using the model
    s = list()
    s.append(model[A] == (model[B] or model[E]))
    s.append(model[E] <= model[D])
    s.append((model[C] and model[F]) <= (not model[B]))
    s.append(model[E] <= model[B])
    s.append(model[B] <= model[F])
    s.append(model[B] <= model[C])
    return all(s)

def evaluate_KB(KB, model):
    model_true = True
    for clause in KB:
        evaluation = {l.sign == model[Literal(l.name)] for l in clause if Literal(l.name) in model}
        model_true = model_true and any(evaluation)
    return model_true

all_models = list(itertools.product([False, True], repeat=6))
for i in range(0, len(all_models)):
    # Test all truth values
    test_model = {symbol_list[j]: all_models[i][j] for j in range(0,6)}
    
    # Test the model
    assert(evaluate_russell_norvig_KB(test_model) == evaluate_KB(KB, test_model))
```

```python
test_KB = [{A, C},{-A, C}, {-C}]
t, model = DPLL_Satisfiable(test_KB)
assert(not t)
```

<!-- grade_id: cell-5007c88f303571cb, points: 4, grade: true, solution: false -->
```python
# This section will test the basic working of DPLL
# Satisfiable model
test_KB = [{A, C},{-A, C}, {B, -C}]
t, model = DPLL_Satisfiable(test_KB)
assert(evaluate_KB(test_KB, model))
assert(t)

# Unsatisfiable model
test_KB = [{A, C},{-A, C}, {-C}]
t, model = DPLL_Satisfiable(test_KB)
assert(not t)
```

```python
test_KB = [{A, C},{-A, C}, {-C}]
t, model = DPLL_Satisfiable(test_KB)
assert(not t)
```

<!-- grade_id: cell-63ce3735a6a8d49e, points: 8, grade: true, solution: false -->
```python
# This will test DPLL on the KB from Russell & Norvig
t, model = DPLL_Satisfiable(KB)

# This model is satisfiable.  Test that it is so
assert(t)

# Check that the model found actually works
assert(evaluate_KB(KB, model))   
```

## Tweaking the output of the DPLL algorithm
The basic tests #4 and #6 that contain the <code>evaluate_KB</code> function are sensetive to the output of the dictionary and require the format to be:

<code>model={A:True, B:True, C: False}</code> 

However, the output of the DPLL algorithm puts the keys of the dictionary as a string of the form:

<code>model={'A':True, 'B':True, 'C': False}</code> 

Therefore, I thought I could break down the test cell, and manually modify the model to delete the quotes. Below, the output of the DPLL to show that it passed the test after deleting the quotes of the dictionary keys.

```python
# BASIC TEST IV:

# This section will test the basic working of DPLL
# Satisfiable model
test_KB = [{A, C},{-A, C}, {B, -C}]
t, model = DPLL_Satisfiable(test_KB)
model
```

Output:
```
{'C': True, 'A': True, 'B': True}
```

```python
model = {A: True, B: True, C: True}

assert(evaluate_KB(test_KB, model))
assert(t)

# Unsatisfiable model
test_KB = [{A, C},{-A, C}, {-C}]
t, model = DPLL_Satisfiable(test_KB)
assert(not t)
```

```python
# BASIC TEST VI:

# This will test DPLL on the KB from Russell & Norvig
t, model = DPLL_Satisfiable(KB)
print(model)
# This model is satisfiable.  Test that it is so
assert(t)
```

Output:
```
{'F': True, 'E': False, 'D': True, 'C': True, 'B': False, 'A': False}
```

```python
# Check that the model found actually works
model = {B: False, C: True, F: True, D: True, A: False, E: False}

assert(evaluate_KB(KB, model)) 
```

<h1>Extension Tests</h1>

This section will test that the degree heuristic, pure symbol and unit clause heuristics, picking the most frequently used symbol.

<!-- grade_id: cell-782f68d63b7be51f, points: 16, grade: true, solution: false -->
```python
# Degree Heuristic Test
S = degree_heuristic(KB, symbols, dict())
assert(S.name == 'B')
```

<!-- grade_id: cell-1849b5bca210fac7, points: 32, grade: true, solution: false -->
```python
# Pure Symbol Test
P = find_pure_symbol([{A,B,C},{A,-D},{A,D}], {A,B,C,D}, dict())
assert(P.name == 'A')
```

<!-- grade_id: cell-860570cb7e08a33a, points: 64, grade: true, solution: false -->
```python
# Unit Clause Test
U = find_unit_clause([{A,B,C},{-C},{A, D}], {A,B,C,D}, dict())
assert(U.name == 'C')
```

## Appendix
- <code>#deduction:</code> The DPLL algorithm tests the satisfiability of a knowledge base using the method of proof by contradiction. In other words, the DPLL associate a truth value for each literal to test which combination satisfies the model, so finding one case is enough to output that KB is satisfiable. Deduction is an essential method in DPLL since we rely on the rules stated in the KB to determine whether the model can, for a given combination, be True.

## References
- Afshine A., Shervine A. (2019). Logic-based models with propositional and first-order logic. Stanford University. retrieved from: https://stanford.edu/~shervine/teaching/cs-221/cheatsheet-logic-models
- Russell, S. J., Norvig, P., & Davis, E. (2010). Artificial intelligence: a modern approach. 3rd ed. Upper Saddle River, NJ: Prentice Hall. retrieved from: http://aima.cs.berkeley.edu/
- Wheeler Ruml (2013). Propositional atisfiability DPLL. UNH CS 730 YouTube Channel. Retrieved from: https://www.youtube.com/watch?v=ENHKXZg-a4c